```python
NAME = ""
COLLABORATORS = ""
```

# Instructions

1. Make sure you have filled out your "NAME" and "COLLABORATORS" (if any) in the previous cell.

2. You should complete all code/markdown cells that state "YOUR CODE HERE" or "YOUR ANSWER HERE". 
   
3. Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

4. Partial credit can be obtained if your solution approach is clear and the documented within comments in the implementation.

5. You should follow good coding practices. Your code should use type hints, be robust against invalid inputs, and you should also write a few test cases to check for correctness particularly including edge cases.  


## Problem 1

Write a python function which computes third power of a number if the number is odd, and the square of the number if the number is even.

<!-- grade_id: cell-14b79c9cb73bbf40, points: 0, grade: false, solution: true -->
```python
def parity_power(x: int) -> int:
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# provide test cases

# YOUR CODE HERE
assert parity_power(1) == 1
```

## Problem 2

An important part of this course is to complete an ML based project over the course of the semester.  We are going to start right now!

Read the description of the project uploaded on canvas and write down 2 potential ideas for your project.  

<!-- grade_id: cell-e119bc622f6ff41b, points: 10, grade: true, solution: true -->
YOUR ANSWER HERE

## Problem 3

This problem will test some of your programming skills and data wrangling ability.  

We will start with a dataset from 'https://github.com/JeffSackmann/tennis_atp' which covers information on Tennis matches from the past few decades.  

This dataset is provided by Jeff Sackman.

Write a function *download_data_by_year* which takes the year as an input and returns a pandas dataframe with the data from the file of the form `atp_matches_{year}.csv`.  Implement a cache so that we do not download the data if it is already in the cache. 

**Hint**: The url you need to use to access the files in github is https://raw.githubusercontent.com/JeffSackmann/tennis_atp/master/atp_matches_{year}.csv

<!-- grade_id: cell-9954e13bab216e43, points: 0, grade: false, solution: true -->
```python
import pandas as pd
def download_data_by_year(year: int) -> pd.DataFrame:
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-bcd973ba4d369a21, points: 10, grade: true, solution: false -->
```python
df = download_data_by_year(2017)
assert df.shape == (2911, 49)
```

## Problem 4
Download the data for years 2000-2020 (inclusive).  Compute the average number of matches played per year on each surface for each month. 

Return your solution as a dictionary where the key is the (surface, month) pair as a tuple and the value is the average.

**Hint**: You may have a case where no matches were played on a given surface in a month.  This should factor into your calculation as a zero.

<!-- grade_id: cell-be49d8719f86849f, points: 0, grade: false, solution: true -->
```python
def get_statistics():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8cb61f127b11a709, points: 10, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose
res = get_statistics()
assert_allclose(res[('Carpet', 1)], 9)
```

# Problem 5

Determine the BEST 5 players of all time.  There is not a definative answer here, this is your chance to show your creativity.  Please also explain how you arrived at your rankings.  You are free to use web resources to support your answer, but you MUST cite them as you use them.

You should answer questions like:

1. How did you define BEST?
2. Where do you believe your analysis is flawed?
3. What could you do to improve your analysis?

```python
# YOUR CODE HERE
```

YOUR ANSWER HERE